# Project Euler: Problem 52

https://projecteuler.net/problem=52

It can be seen that the number, $125874$, and its double, $251748$, contain exactly the same digits, but in a different order.

Find the smallest positive integer, $x$, such that $2x$, $3x$, $4x$, $5x$, and $6x$, contain the same digits.

First, write a function `same_digits(x,y)` that returns `True` if two integers `x` and `y` have the exact same set of digits and multiplicities and `False` if they have different digits.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
def same_digits(x, y):
    """Do the integers x and y have the same digits, regardless of order."""
    x=str(x)
    y=str(y)             #turn the inputs into string
    b=[]                
    k=[]
    for i in range(len(x)):       # add each individual digit as elements into our empty lists
        b.append(x[i])
    for i in range(len(y)):
        k.append(y[i])
    if len(b) != len(k):         # ensures that both x and y inputs have the same lengths
        return False
    else:
        for q in range(len(b)):       # for each element in list b
            for w in range(len(k)):
                if q==k[w]:           # if that element is equal to the one in k...
                    return True
                else:
                    return False
        
            
            
```

```python

```

<!-- grade_id: projecteuler52a, points: 4, grade: true, solution: false -->
```python
assert same_digits('132', '321')
assert not same_digits('123', '3')
assert not same_digits('456', '0987654321')
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-61-4d2369f19b4d> in <module>()
----> 1 assert same_digits('132', '321')
      2 assert not same_digits('123', '3')
      3 assert not same_digits('456', '0987654321')

AssertionError:
```

Now use the `same_digits` function to solve this Euler problem. As you work on this problem, be careful to debug and test your code on small integers before trying it on the full search.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: projecteuler52b, points: 6, grade: true, solution: false -->
```python
assert True # leave this cell to grade the solution
```