# Integration Exercise 3

## Imports

```python
%matplotlib inline
import matplotlib.pyplot as plt
import numpy as np
from scipy import integrate
```

## 2d polar integration

The 2d polar integral of a scalar function $f(r, \theta)$ is defined as:

$$ I(r_{max}) = \int_0^{r_{max}} \int_0^{2\pi} f(r, \theta) r d\theta dr $$

Write a function `integrate_polar(f, rmax)` that performs this integral numerically using `scipy.integrate.dblquad`.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
def integrate_polar(f, rmax):
    """Integrate the function f(r, theta) over r=[0,rmax], theta=[0,2*np.pi]"""
    
    I,e = integrate.dblquad(integrate_polar, 0 , 2*np.pi,0,rmax)
    return I
```

<!-- grade_id: integrationex04a, points: 10, grade: true, solution: false -->
```python
assert np.allclose(integrate_polar(lambda r,t: 1, 1.0), np.pi)
assert np.allclose(integrate_polar(lambda r, t: np.exp(-r)*(np.cos(t)**2), np.inf), np.pi)
```

Output:
```
/usr/local/lib/python3.4/dist-packages/scipy/integrate/quadpack.py in _quad(func, a, b, args, full_output, epsabs, epsrel, limit, points)
    374     if points is None:
    375         if infbounds == 0:
--> 376             return _quadpack._qagse(func,a,b,args,full_output,epsabs,epsrel,limit)
    377         else:
    378             return _quadpack._qagie(func,bound,infbounds,args,full_output,epsabs,epsrel,limit)

/usr/local/lib/python3.4/dist-packages/scipy/integrate/quadpack.py in _infunc(x, func, gfun, hfun, more_args)
    441 
    442 def _infunc(x,func,gfun,hfun,more_args):
--> 443     a = gfun(x)
    444     b = hfun(x)
    445     myargs = (x,) + more_args

TypeError: 'int' object is not callable
```

```python

```